<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 3002782  -->
  <question type="stack">
    <name>
      <text>Operaciones combinadas naturales</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Realiza las siguientes operaciones:</p>
<p>a) \({@n1@}-({#n2#}-{#n3#})=\)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>b) \({@n4@}+{#n5#} \times {#n6#}=\)[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
<p>c) \(10 \times {#n7#} - 10 \times {@n8@}=\)[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2021120900</text>
    </stackversion>
    <questionvariables>
      <text>/* apdo a): n1-(n2-n3) */
n1: rand_with_step(10,15,1);
n2: rand_with_step(6,9,1);
n3: rand_with_step(2,4,1);
solA: n1-(n2-n3);

/* apdo b): n4-(n5-n6) */
numsB: rand_selection([2,3,4,5,6,7,8,9],3);
n4: numsB[1];
n5: numsB[2];
n6: numsB[3];
solB: n4+n5*n6;

/* apdo c): 10*n7-10*n8 */
numsC: sort(rand_selection([2,3,4,5,6,7,8,9],2));
n7: numsC[2];
n8: numsC[1];
solC: 10*n7-10*n8;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text>\({@n1@}-({#n2#}-{#n3#})=   {@n4@}+{#n5#} \times {#n6#}=  10 \times {#n7#} - 10 \times {@n8@}=\)</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> La teva resposta és parcialment correcta.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>solA</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>solB</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>solC</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>solA</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>n1-n2-n3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<span>Primero se calcula el paréntesis. </span>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans2</sans>
        <tans>solB</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>(n4+n5)*n6</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<span>Las multiplicaciones se hacen antes que las sumas.</span>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans3</sans>
        <tans>solC</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1903104628</deployedseed>
    <deployedseed>1753355995</deployedseed>
    <deployedseed>313912397</deployedseed>
    <deployedseed>1743160380</deployedseed>
    <deployedseed>1852991320</deployedseed>
    <deployedseed>297893557</deployedseed>
    <deployedseed>1351403265</deployedseed>
    <deployedseed>1072187680</deployedseed>
    <deployedseed>1846761465</deployedseed>
    <deployedseed>1149653002</deployedseed>
    <deployedseed>168783382</deployedseed>
    <deployedseed>344313219</deployedseed>
    <deployedseed>219308297</deployedseed>
    <deployedseed>1683648184</deployedseed>
    <deployedseed>1706983417</deployedseed>
    <deployedseed>737966756</deployedseed>
    <deployedseed>923841366</deployedseed>
    <deployedseed>671167411</deployedseed>
    <deployedseed>285321139</deployedseed>
    <deployedseed>1889553566</deployedseed>
    <deployedseed>998816698</deployedseed>
    <deployedseed>701514082</deployedseed>
    <deployedseed>1751739062</deployedseed>
    <deployedseed>1105793283</deployedseed>
    <deployedseed>1620337645</deployedseed>
    <deployedseed>975321352</deployedseed>
    <deployedseed>1992196420</deployedseed>
    <deployedseed>2009789210</deployedseed>
    <deployedseed>85877274</deployedseed>
    <deployedseed>1091726778</deployedseed>
    <deployedseed>1249144887</deployedseed>
    <deployedseed>1909603621</deployedseed>
    <deployedseed>1599589825</deployedseed>
    <deployedseed>159595156</deployedseed>
    <deployedseed>65261477</deployedseed>
    <deployedseed>995841142</deployedseed>
    <deployedseed>1475104092</deployedseed>
    <deployedseed>957531873</deployedseed>
    <deployedseed>1792362860</deployedseed>
    <deployedseed>467608040</deployedseed>
    <deployedseed>155335830</deployedseed>
    <deployedseed>1837130092</deployedseed>
    <deployedseed>1489279817</deployedseed>
    <deployedseed>1012335842</deployedseed>
    <deployedseed>239797452</deployedseed>
    <deployedseed>902660054</deployedseed>
    <deployedseed>1288117631</deployedseed>
    <deployedseed>1986812860</deployedseed>
    <deployedseed>415554389</deployedseed>
    <deployedseed>6860876</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>solA</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>solB</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>solC</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>